# Lab 9
Systems of linear differential equations can be solved using matrix calculations. We have covered numerical matrix calculations in this course already, but today we will recap and cover some advanced operations.

In the exercises you will employ some of these skills to use the computer to find analytic solutions to coupled differential equations.

## Setup
We adopt a common custom and import `numpy` as `np`. This is because we will be using a lot of NumPy functions today so this keeps our imports neat.

```python
import numpy as np
from numpy import isclose, allclose, array, exp, linspace, outer
from numpy.linalg import det, inv, eig, solve, eigvals
from scipy.integrate import odeint
import pandas as pd
import seaborn as sns
```

## Entering matrices
Seeing this one more time doesn't hurt. We can enter the matrix
\begin{align*}
A = \left(\begin{array}{3} 1&2&0 \\ 2&1&0 \\ 0&1&2\end{array}\right)
\end{align*}
in the following fashion.

```python
A = np.array([[1, 2, 0],
           [2, 1, 0],
           [0, 1, 2]])
print(A)
```

Which element do we select with `A[0, 2]`?

```python
A[0, 2]
```

Which column is selected by `A[:, 2]`?

```python
A[:, 2]
```

Notice that NumPy doesn't distinguish between column and row vectors (unless you really force it). It usually figures out what you need from the context.

Which row is selected by `A[0, :]`?

```python
A[0, :]
```

Recall that addition is elementwise. To add the first two columns of `A` and label them `x` write

```python
x = A[:, 0] + A[:, 1]
print(x)
```

## Special matrices and operations

`eye` returns an identity matrix.

```python
np.eye(3)
```

`zeros` returns a zero array. Note that it needs both dimensions as it makes no assumptions about the shape you want.

```python
np.zeros((4, 4))
```

`diag` takes a vector and returns a matrix with that vector as its diagonal.

```python
np.diag([1, 2, 3])
```

We can transpose a matrix with its member property `T`.

```python
A.T
```

We can take the determinant using `det`.

```python
det(A)
```

`inv` returns the inverse matrix.

```python
inv(A)
```

Recall that the `@` operator performs matrix multiplication.

```python
inv(A) @ A
```

Note that `inv` and `det` come from the `numpy.linalg` module.
## Solving a system of linear equations
`solve`, also from `numpy.linalg`, solves matrix equations of the form
\begin{align*}
Ax=b
\end{align*}
for $x$, given $A$ and $b$. For example, take `A` as above and imagine
\begin{align*}
b=\left(\begin{array}{c}1\\2\\3\end{array}\right).
\end{align*}


```python
b = np.array([1, 2, 3])
solve(A, b)
```

`inv(A) @ b` gives the same result, of course.

```python
inv(A) @ b
```

## Finding eigenvalues and eigenvectors
You have seen before that you can obtain the eigenvalues of a matrix using `eigvals`.

```python
eigvals(A)
```

It is rare that you want the eigenvalues in isolation, however. `eig` gives you the eigenvalues and the eigenvectors.

```python
v, W = eig(A)
print(v)
print(W)
```

The eigenvectors are now in the columns of `W`. For instance, we can verify the eigenvalue equation for the second eigenvalue and eigenvector.

```python
print('A w_1:', A.dot(W[:, 1]))
print('v_1 w_1:', v[1]*W[:, 1])
```

## Linear independence of a set of vectors
To test for linear independence of a set of vectors $v_1, v_2, v_3, \ldots, v_4$, set
\begin{align*}
c_1v_1+c_2v_2+c_3v_3+\ldots+c_nv_n = 0,
\end{align*}
then show that the only solution to these equations is where $c_1=c_2=c_3=\ldots=c_n=0$.

If the $U=\left(\begin{array}{ccccc}v_1&v_2&v_3&\ldots&v_n\end{array}\right)$ is a matrix and $c$ is the (column) vector of the coefficients $c_i$, we can write the equations in the matrix form $Uc=0$.

If the determinant of $U$ is non-zero, $U$ has an inverse (is non-singular) and the solution to the equations is $c=U^{-1}0=0$. Therefore, the $c_i$s are all 0 and the vectors are linearly independent.

For example, we can show that our eigenvectors from the last section are linearly independent.

```python
det(W)
```

In practice, this approach is fraught with danger, because `det(W)` is a floating point result so can be almost zero when it would be exactly zero if we performed the calculations exactly. However, if you find that the determinant is nearly zero, then you can assume that bad things will happen when you try to invert the matrix.

## Linear combinations of vectors
To write a vector $r$ as a linear combination of the vectors $v_1, v_2, v_3, \ldots, v_4$, set
\begin{align*}
r = a_1v_1 + a_2v_2 + a_3v_3 + \ldots + a_nv_n,
\end{align*}
and solve these euqations to find the coefficients $a_1$, $a_2$, $a_3$, $\ldots$, $a_n$.

In matrix form
\begin{align*}
Ua=r,
\end{align*}
where $a$ is the (column) vector of the coefficients $a_i$. Now we can use the command `solve(U, r)` to find the vector $a$ and hence the $a_i$s.

# Exercises

We will use the computer to solve coupled differential equations analytically and numerically.

Recall that the system of differential equations 
\begin{align*}
\frac{\mathrm{d}\mathbf{y}}{\mathrm{d}t} = \mathbf{A}\mathbf{y}(t)
\end{align*}
where $\mathrm{A}$ is a $2\times 2$ matrix with distinct eigenvalues, has the solution
\begin{align*}
\mathbf{y}(t) = a\mathrm{e}^{\lambda_1}\mathbf{v}_1 + b\mathrm{e}^{\lambda_2}\mathbf{v}_2,
\end{align*}
where $\lambda_1$ and $\lambda_2$ are the eigenvalues of $\mathbf{A}$ and $\mathbf{v}_1$ and $\mathbf{v}_2$ are the corresponding eigenvectors.

Take
\begin{align*}
\mathbf{A} = \left(\begin{array}{cc}-1.5&0.5 \\ 1&-1\end{array}\right).
\end{align*}
In the cell below, use `numpy.linalg.eig` to find the eigenvalues and eigenvectors of $\mathbf{A}$.

<!-- grade_id: cell-c8203086265517f7, points: 1, grade: true, solution: true -->
```python

```

If $\mathbf{y}(0)=\left(\begin{array}{c}5\\4\end{array}\right)$, use `numpy.linalg.solve` to calculate $a$ and $b$.

<!-- grade_id: cell-40bb20b105652d9a, points: 1, grade: true, solution: true -->
```python

```

In the cell below, write a function that takes a NumPy array of $t$ values as an input and outputs the corresponding array of $\mathbf{y}$ values.

_Hint:_ you could iterate through the $t$ values to calculate each value of $\mathbf{y}$ using a for loop, or you could use expressions like `a*outer(exp(v[0]*t), W[:,0])` to calculate what is known as the [outer product](https://docs.scipy.org/doc/numpy/reference/generated/numpy.outer.html).

<!-- grade_id: cell-75073d4d28ba65c1, points: 0, grade: false, solution: true -->
```python
def y(t):
    ### YOUR IMPLEMENTATION GOES HERE
```

<!-- grade_id: cell-beb811ce1aa73b0e, points: 1, grade: true, solution: false -->
```python
assert allclose(y(array([0, 2.5, 5])), [[5.        , 4.        ],
                                        [0.87299028, 1.70555289],
                                        [0.2463458 , 0.49241919]])
assert allclose(y(array([10])), [0.02021385, 0.04042768])
assert allclose(y(array([20])), [0.0001362, 0.0002724])
```

Below we will calculate the solution numerically using `odeint`, so in the this cell implement a function to calculate $\mathrm{d}\mathbf{y}/\mathrm{d}t$ given $\mathbf{y}$ and $t$.

<!-- grade_id: cell-4cfbde3758e27da5, points: 0, grade: false, solution: true -->
```python
def difeq(y, t):
    ### YOUR IMPLEMENTATION GOES HERE
```

<!-- grade_id: cell-50efc729305ba8c7, points: 1, grade: true, solution: false -->
```python
assert allclose(difeq([5, 4], 0), [-5.5,  1. ])
assert allclose(difeq([0.87299028, 1.70555289], 2.5), [-0.45670898, -0.83256261])
assert allclose(difeq([0.2463458 , 0.49241919], 5), [-0.1233091 , -0.24607339])
```

Finally, calculate the numerical solution using `odeint` and plot it with the analytic solution on same graph over the interval $0\leq t\leq5$.

<!-- grade_id: cell-275d54fac40d5aa1, points: 2, grade: true, solution: true -->
```python

```