# Lab 6
## Introduction
This lab covers some familiar ground then introduces numerical root finding, matrix manipulation, and curve fitting using NumPy and SciPy.

Numerical root finding could be useful for estimating a single critical model parameter.

Matrix calculations will be critical for solving systems of differential equations. Being able to handle matrices numerically is also a useful life skill. When your first-year linear algebra is a distant memory and your continued employment is contingent on you finding an eigenvalue, it could help you find your way out of a dark place.

Curve fitting will be useful if you want to figure out the parameters of a model given some empirial data.

We will cover

1. Basic Maths
1. Plotting Functions
1. Finding Roots of Functions
1. Finding Eigenvalues
1. Curve Fitting

Hopefully the first two sections are old hat to you. They are included here as a summary.

### Setup
As always, start with some imports.

```python
from numpy import linspace, exp, sqrt, cos, sin, pi, array, random, meshgrid, sqrt
from numpy.linalg import eigvals
from numpy.testing import assert_almost_equal
from scipy.optimize import fsolve, curve_fit
from scipy.integrate import odeint
from plotly import graph_objs as go
from plotly.figure_factory import create_quiver
```

### 1. Basic Maths

As you know, Python works just fine as a calculator.

```python
2 + 3
```

```python
7 - 5
```

```python
34 * 212
```

```python
1234 / 5678
```

```python
3**7
```

```python
sqrt(2)
```

```python
cos(pi)
```

```python
sin(pi)
```

Note that the last answer is $1.2246\ldots\times 10^{-16}$, which is very very small and should remind you that non-integer calculations always come with a tiny bit of error.

### 2. Plotting Functions
Plotting functions is a very important aspect of mathematics. It gives a very quick and visual way of understanding the behaviour of functions. When we plot a function, say $f(x) = \sin x$, we first must define the range of $x$. Let’s say we want to plot the curve $y = \sin x$ in the range $−5 \leq x \leq 5$. We define the range of x by typing the following command.

```python
x = linspace(-5, 5, 101)
```

This creates a vector `x` whose elements are from −5 to 5 in steps of 0.1. The vector `x` has 101 elements in the vector. (Note that `x` is actually a NumPy array. NumPy calls matrices and vectors "arrays".) You can see what is in `x` like so:

```python
print(x)
```

```python
y = sin(x)
fig = go.Figure()
fig.add_trace(go.Scatter(x=x, y=y))
fig.show('png')
```

Note that you can download the plot (so that you can include it in your assignment) by pressing the camera icon.

You can decorate your plot. Remember that `go.Figure()` created the figure object in memory, to which you added a scatter plot before calling `fig.show()` to display it in your notebook. The same figure is still in memory, so the next command just updates it.

```python
fig.layout.update(dict(title='A plot of y=sin(x)',
                       xaxis=dict(title='x-axis label'),
                       yaxis=dict(title='y-axis label')))
fig.show('png')
```

Hence for basic plotting all you need is to do is: (i) define the range; (ii) define the function you want to plot; and (iii) type in the plot command. It is as easy as that! You can then add the other bits to “beautify” your plot, like grids, axis labels, title, etc.

Let’s try plotting another curve $y = x2 − x \sin(3x) − x\mathrm{e}^x$ for $−3 \leq x \leq 3$. Note that when you raise $x$ to the power of 2, multiply it by other arrays, or apply NumPy functions like `sin` and `exp`, NumPy assumes that you want to perform _elementwise_ operations.

```python
x = linspace(-3, 3, 61)
y = x**2 - x*sin(3*x) - x*exp(x)
fig = go.Figure()
fig.add_trace(go.Scatter(x=x, y=y))
fig.show('png')
```

### 3. Finding Roots of Functions

In many circumstances one would want to find the value of x when the function $f(x) = 0$, i.e., the roots of $f(x)$. Mathematicians, statisticians, and computer scientists have quite a lot to say about root finding but in many cases you can get away with calling a SciPy function.

The first thing you need is a “guess” for the root. A good way of getting an appropriate guess is to first plot the function, and then zoom in, or simply by sight.

Say for example, we would like to find a root of the function $f(x) = x−\sin x+1$. So we first plot the curve $y = f(x)$ to see “roughly” where the curve “cuts” the x-axis. Remember the three lines of code that will plot this curve?

```python
x = linspace(-3, 3, 61)
y = x - sin(x) + 1
fig = go.Figure()
fig.add_trace(go.Scatter(x=x, y=y))
fig.show('png')
```

So from our naked eye, we can guess that the value of the root of the function is around $x = −2$. So now we type in the following commands to get the a more accurate answer.

```python
def func(x):
    return x - sin(x) + 1

fsolve(func, -2)
```

If you want to be fancy, you can do it all on one line using a _lambda_ function. Lambda functions exist in several programming languages and are convenient for defining a function you only intend to use once.

```python
fsolve(lambda x: x - sin(x) + 1, -2)
```

### 4. Finding Eigenvalues
Eigenvalues provide very important information regarding the behaviour of various systems. In the next few weeks we will be learning how we can use matrix methods and eigenvalues to solve linear systems of differential equations. Once we go beyond $2 \times 2$ matrices, finding eigenvalues by hand is a pain and using Python easy. How do we find the eigenvalues of matrices using Python? It is disgustingly simple! Read on and find out.

Say we would like to find the eigenvalues of matrix A whose elements are
\begin{align*}
\left(\begin{array}{2}
1&2\\
3&4
\end{array}\right)
\end{align*}

We first need to define matrix `A` as a NumPy array and we do this by typing the following command

```python
A = array([[1, 2], [3, 4]])
```

which gives

```python
print(A)
```

To obtain the eigenvalues we simply type

```python
eigvals(A)
```

So what? Anyone can find the eigenvalues of a $2 \times 2$ matrix, I hear you say!

Let’s now find the eigenvalues of a $3 \times 3$ matrix `B` whose elements are
\begin{align*}
\left(\begin{array}{2}
1&2&3\\
4&5&6\\
7&8&9
\end{array}\right)
\end{align*}


Imagine having to find the eigenvalues of this $3 \times 3$ matrix by hand. In Python, as before, we need to first define the matrix, and then use the `eigvals` command as follows.

```python
B = array([[1, 2, 3],
           [4, 5, 6],
           [7, 8, 9]])
eigvals(B)
```

Note that one of the eigenvalues is zero.

How easy was that! You can do this for any $n \times n$ matrix. Why don’t you make up your own $6 \times 6$ matrix and find the eigenvalues of that matrix. It’s the same two commands you need. One to define the matrix, and then the `eigvals` command.

```python

```

If you would like more practice (and information) on using NumPy arrays, please check out [this lesson](https://swcarpentry.github.io/python-novice-inflammation/02-numpy/index.html) from Software Carpentary.

### 5. Curve Fitting
The following material is just the example from the SciPy documentation that is available [here](https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.curve_fit.html).

Start by defining a model that we want to fit.

```python
def func(x, a, b, c):
    return a * exp(-b * x) + c
```

Define the data to be fit with some noise:

```python
xdata = linspace(0, 4, 50)
y = func(xdata, 2.5, 1.3, 0.5)
random.seed(1729)
y_noise = 0.2 * random.normal(size=xdata.size)
ydata = y + y_noise

fig = go.Figure()
fig.add_trace(go.Scatter(x=xdata, y=ydata, name='data'))
fig.show('png')
```

Fit for the parameters a, b, c of the function `func`:

```python
popt, pcov = curve_fit(func, xdata, ydata)
print('a, b, c = ', popt)
a, b, c = popt
yfitted = func(xdata, a, b, c)

fig.add_trace(go.Scatter(x=xdata, y=yfitted, name='fitted'))
fig.show('png')
```

How close were the fitted parameters to the actual parameters?

## Exercises

### Slope field and DE solution plot

Plot on the one figure the slopefield for the DE
\begin{align*}
\frac{\mathrm{d} y}{\mathrm{d} x} = 2.5y (1 − y)\qquad y(0) = 0.5,
\end{align*}
and the solutions $y(x)$ with $y(0) = 0.2$, $y(0) = 0.5$ and $y(0) = 0.8$.

Start by writing down a new definition for `diff_eq` below. Do not change the function's name or inputs.

<!-- grade_id: cell-05cc4f7824ab2d84, points: 0, grade: false, solution: true -->
```python
def diff_eq(y, x):
    ### your implementation goes here
```

If you have implemented `diff_eq` correctly, the following should print "nice job".

<!-- grade_id: cell-0a0fa9099e30078d, points: 1, grade: true, solution: false -->
```python
assert_almost_equal(diff_eq(0.4, 0), 0.6)
assert_almost_equal(diff_eq(0.4, 10), 0.6)
assert_almost_equal(diff_eq(1, 0), 0)
print("nice job")
```

Now create your graph. Note that you will have to redefine `S` (from the lab). You can do that using your new definition for `diff_eq` or by writing out the RHS of the equation again.

You will also have to change your definition of the meshgrid for the slopefield and the domain and initial values in the `odeint` commands. You want about 21 steps in the x and y ranges in meshgrid. If you change the scaling factor from 0.3 to 0.04 in `create_quiver`, you will get a better slope field.

Create the plot for the region $0 < x < 1$ and $0 < y < 1$.

<!-- grade_id: cell-8945b9507fff370f, points: 2, grade: true, solution: true -->
```python

```

## Solution at a point
Write code in the next cell so that the last line is the single number for $y(1)$ if $y(0)=0.8$.

<!-- grade_id: cell-962d55b6bbeb85ad, points: 0, grade: false, solution: true -->
```python

```