# Lab 11
This week we will take a closer look at how the long-run behaviour of a system is decided by its eigenvalues (and eigenvectors).
## Setup

```python
from numpy import (real, imag, isclose, allclose, 
                   array, exp, linspace, cos, sin,
                   meshgrid, sqrt, round, arange, ones)
from numpy.linalg import eig, solve
from scipy.integrate import odeint
from scipy.linalg import expm, logm
import plotly.io
from plotly import graph_objs as go
from plotly.figure_factory import create_quiver
```

The vector $\mathbf{x}$ gives the probability that a nucleotide is in a particular state (A, C, G, or T) in the human gene ENSG00000069248. Over the last 160 million years, $\mathbf{x}$ evolved according to the equation
\begin{align*}
\mathbf{x}_{t+1}=\left[\begin{matrix}
0.9938& 0.0024& 0.0019& 0.0001\\
0.0005& 0.9936& 0.0006& 0.0033\\
0.0045& 0.0006& 0.9969& 0.0003\\
0.0012& 0.0034& 0.0006& 0.9963\end{matrix}\right]\mathbf{x}_t,
\end{align*}
where each time increment ($t$ to $t+1$) corresponds to 10 million years.

For your convenience, I have copied the above matrix into a NumPy array:

```python
P = array([[9.938e-01, 2.400e-03, 1.900e-03, 1.000e-04],
           [5.000e-04, 9.936e-01, 6.000e-04, 3.300e-03],
           [4.500e-03, 6.000e-04, 9.969e-01, 3.000e-04],
           [1.200e-03, 3.400e-03, 6.000e-04, 9.963e-01]])
```

1. Given that
\begin{align*}
\mathbf{x}_0=\left[\begin{matrix}
0.3082\\0.1822\\0.1663\\0.3433
\end{matrix}\right],
\end{align*}
and that $t=16$ today, plot $\mathbf{x}_t$ for the last 160 million years. Plot all four probabilities on the same graph, using markers rather than lines.

<!-- grade_id: cell-acb6f7bb4867ff40, points: 2, grade: true, solution: true -->
```python

```

2. The solution to the equation above takes the form
\begin{align*}
\mathbf{x}_t=a_1\lambda_1^t\mathbf{v}_1 + a_2\lambda_2^t\mathbf{v}_2 + a_3\lambda_3^t\mathbf{v}_3 + a_4\lambda_4^t\mathbf{v}_4,
\end{align*}
where $\lambda_i$ and $v_i$ are the eigenvalues and eigenvectors of $P$, respectively, and $a_i$ are scalars. Find the values of $a_i$ for the initial conditions ($x_0$) given in question 1.

<!-- grade_id: cell-8f8cbc5cedc9e118, points: 1, grade: true, solution: true -->
```python

```

3. Use the formula from Question 2 to write a function that returns values of $x_t$ given an array of times, the $a_i$ coefficients, the eigenvalues, and the eigenvectors. The latter two should be in the form as returned by `numpy.linalg.eig`.

<!-- grade_id: cell-e9e8638bdd4ded30, points: 0, grade: false, solution: true -->
```python
def calc_x(t, a, lam, V):
    ### YOUR IMPLEMENTATION GOES HERE
```

<!-- grade_id: cell-b97906d084fd6c87, points: 1, grade: true, solution: false -->
```python
assert allclose(calc_x(
    arange(4), array([1, 1, 1, 1]), array([1, 0.9, 0.8, 0.6]), 
    (ones((4,4))*[0.1, 0.2, 0.3, 0.4]).T),
                array([[0.4       , 0.8       , 1.2       , 1.6       ],
                       [0.33      , 0.66      , 0.99      , 1.32      ],
                       [0.281     , 0.562     , 0.843     , 1.124     ],
                       [0.2457    , 0.4914    , 0.7371    , 0.9828    ]]))
lam, v = eig(P)
a = solve(v, x0)
assert allclose(calc_x(arange(17), a, lam, v),
                 array([[0.3082    , 0.1822    , 0.1663    , 0.3433    ],
                       [0.30707674, 0.18242069, 0.16738368, 0.34311889],
                       [0.30596301, 0.18263946, 0.16845902, 0.3429385 ],
                       [0.30485874, 0.18285632, 0.1695261 , 0.34275884],
                       [0.30376385, 0.18307129, 0.17058497, 0.34257989],
                       [0.30267826, 0.18328438, 0.17163572, 0.34240165],
                       [0.30160188, 0.1834956 , 0.17267839, 0.34222413],
                       [0.30053465, 0.18370498, 0.17371306, 0.34204731],
                       [0.29947649, 0.18391252, 0.17473979, 0.3418712 ],
                       [0.29842732, 0.18411824, 0.17575865, 0.3416958 ],
                       [0.29738706, 0.18432215, 0.1767697 , 0.34152109],
                       [0.29635565, 0.18452426, 0.17777301, 0.34134709],
                       [0.29533301, 0.18472459, 0.17876863, 0.34117378],
                       [0.29431906, 0.18492315, 0.17975663, 0.34100116],
                       [0.29331373, 0.18511996, 0.18073708, 0.34082923],
                       [0.29231696, 0.18531503, 0.18171002, 0.34065799],
                       [0.29132866, 0.18550837, 0.18267554, 0.34048743]]))
```

4. Recreate your plot from Question 1 using your formula.

<!-- grade_id: cell-4c0b4ec1f2fb55ff, points: 1, grade: true, solution: true -->
```python

```

5. Use your formula to extend the plot out to ten billion years ($t=1000$).

<!-- grade_id: cell-2f3de6a55ae6ca8c, points: 1, grade: true, solution: true -->
```python

```

6. Look for the dominant eigenvalue of $P$ and find the corresponding eigenvector. What are the limiting values of $x$ for large $t$? Do they agree with your plot for Question 5?

<!-- grade_id: cell-d56fc5139095b093, points: 1, grade: true, solution: true -->
```python

```