# Lab 10
Up to now we have visualisted the solutions to coupled differential equations as functions of the independent variable. This week we will look at different visualisations that relate the dependent variables to one another.

## Setup

```python
from numpy import (meshgrid, linspace, pi, array, sqrt, cos, sin,
                   real, imag, isclose, allclose, exp, outer)
from numpy.linalg import eig, solve
from scipy.integrate import odeint
import seaborn as sns
import pandas as pd
import matplotlib.pyplot as plt
```

## Coupled Differential Equations

Take the system of equations
\begin{align*}
x' &= y, \\
y' &= -x.
\end{align*}

### Analytic Solution
In Section 4.1, Question 17 we saw that the solutions to these equations are
\begin{align*}
x(t) &= A \cos t + B \sin t, \\
y(t) &= B \cos t - A \sin t.
\end{align*}

What are $A$ and $B$ for the following initial value problems?

1. $x(0)=0$, $y(0)=1$
1. $x(0)=2$, $y(0)=1$
1. $x(0)=0$, $y(0)=5$

### Plotting the trajectories
To plot trajectories of a system of 2 equations we plot the 
$x$ solution on the horizontal axis and the $y$
solution on the vertical axis.

```python
t = linspace(0, pi/2)
data = pd.DataFrame()
for A, B in [[0, 1], [2, 1], [0, 5]]:
    x = A*cos(t) + B*sin(t)
    y = B*cos(t) - A*sin(t)
    more_points = pd.DataFrame({'x': x, 'y': y, 'initial values': str([A, B])})
    data = pd.concat([data, more_points])
ax = sns.lineplot(data=data, x='x', y='y', hue='initial values', sort=False)
```

What happens if you change `t` to `linspace(0, pi/2)`? Why?

Of course, we could alternatively solve the equations using `odeint`. For that we will need a function for the differential equations:

```python
def difeq(x, t):
    M = array([[0, 1],
               [-1, 0]])
    return M.dot(x)
```

The solutions will be the same:

```python
t = linspace(0, 2*pi)
data = pd.DataFrame()
for x0 in [[0, 1], [2, 1], [0, 5]]:
    x = odeint(difeq, x0, t)
    more_points = pd.DataFrame({'x': x[:,0], 'y': x[:,1], 'initial values': str(x0)})
    data = pd.concat([data, more_points])
ax = sns.lineplot(data=data, x='x', y='y', hue='initial values', sort=False)
```

We can’t tell by looking at these graphs which direction the solution is "moving" in time (because so far we’ve only plotted $x$ and
$y$). We can figure that out using _direction fields_. Direction fields (like a slope field but with vectors instead of line segments) can also be used to graphically study solutions for DEs that can’t be solved analytically.

We can create a direction field in almost exactly the same way as we created a quiver plot in weeks 2 and 3. In this instance, the $x$ component of
each vector is the RHS of the $\mathrm{d}x/\mathrm{d}t$ equation, and the $y$ component of each vector is the RHS of the
$\mathrm{d}y/\mathrm{d}t$ equation. We scale all the lines to have he same length.

Note that we need the DE function, not the analytic solution, for the direction field.

```python
x, y = meshgrid(linspace(-6, 6, 20), linspace(-6, 6, 20))
xy_grid = array([x,y]).transpose((1, 0, 2)) # some numpy magic to vectorise the calculations
dx, dy = difeq(xy_grid, None)
L = sqrt(dx**2 + dy**2)
fig, ax = plt.subplots(figsize=(5, 5))
q = ax.quiver(x, y, dx/L, dy/L, scale=25, color='grey')
data = pd.DataFrame()
for x0 in [[0, 1], [2, 1], [0, 5]]:
    x = odeint(difeq, x0, t)
    more_points = pd.DataFrame({'x': x[:,0], 'y': x[:,1], 'initial values': str(x0)})
    data = pd.concat([data, more_points])
ax = sns.lineplot(data=data, x='x', y='y', hue='initial values', sort=False)
```

What happens if you change `scale` to 2 in `create_quiver`?

Alternatively, we can add the $t$ axis in and create a three dimensional plot.

```python
t = linspace(0, 2*pi)
fig = plt.figure(figsize=(5, 5))
ax = fig.gca(projection='3d')
for A, B in [[0, 1], [2, 1], [0, 5]]:
    x = A*cos(t) + B*sin(t)
    y = B*cos(t) - A*sin(t)
    ax.plot(x, y, t, label=str([A, B]))
```

Can you recover the 2D plot from above by looking down the `z` axis?

# Exercises

We will plot the analytic and numerical solutions to some DEs along with their direction fields.

1. Find the general solution to
\begin{align*}
x' &= -0.5x + 2y\\
y' &= -2x -0.5y
\end{align*}
You can use the following cell to, for instance, calculate the eigenvalues and eigenvectors, if you would like.

<!-- grade_id: cell-c2b085d8dc989c38, points: 0, grade: true, solution: true -->
```python

```

2. Plot the analytic solutions for the following initial value problems for $0\leq t\leq 8$.
    1. $x(0)=2$, $y(0)=1.5$
    2. $x(0)=1$, $y(0)=-3$
    3. $x(0)=-2$, $y(0)=-2.5$

<!-- grade_id: cell-fb85351c79c8429c, points: 2, grade: true, solution: true -->
```python

```

3. Define the derivative function that you will use to calculate the direction field and the numerical solution.

<!-- grade_id: cell-cac3f5ddcb4d23f2, points: 0, grade: false, solution: true -->
```python
def difeq(x, t):
    ### YOUR IMPLEMENTATION GOES HERE
```

<!-- grade_id: cell-808ec730fc3ab81b, points: 1, grade: true, solution: false -->
```python
assert allclose(difeq([2, 1.5], 0), [2, -4.75])
assert allclose(difeq([1, -3], 0), [-6.5, -0.5])
assert allclose(difeq([-2, -2.5], 0), [-4, 5.25])
```

3. Plot the direction field and the analytic solutions from question 2 on the same graph. Plot the ranges $-3\leq x \leq 3$ and $-3\leq y \leq 2$.

<!-- grade_id: cell-4cb3140ef760a441, points: 2, grade: true, solution: true -->
```python

```

4. Recreate the plot from question 3 using `odeint` rather than the numerical solution.

<!-- grade_id: cell-df67ab13fb30b6f9, points: 1, grade: true, solution: true -->
```python

```

5. Use a for loop to plot the solutions for the initial value problems where $y(0)=2$ and 21 equally spaced values of $x(0)$ between -2 and 2. Do not plot the direction field.


<!-- grade_id: cell-6871cf9c0d7906ee, points: 2, grade: true, solution: true -->
```python

```

6. Recreate the plot from question 5 in three dimensions, where the $z$ values correspond to the independent variable. Use lines rather than markers.

<!-- grade_id: cell-7d5627fa91452a15, points: 1, grade: true, solution: true -->
```python

```